---
output: 
  html_document:
    includes:
      in_header: test1.html
runtime: shiny_prerendered
description: "datviz1"
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
knitr::opts_chunk$set(echo = FALSE)
tutorial_options(exercise.eval = FALSE)
knitr::opts_chunk$set(exercise.checker = gradethis::grade_learnr)

```

```{css, echo=FALSE}
p {
  font-size: 16.5px;
  font-family: "Merriweather", sans-serif;
}

h5 {
  font-size: 16.5px;
  font-family: "Lato", sans-serif;
}

```

```{r prepare-aovsetup, warning=FALSE, message=FALSE}
library(ggplot2)

midus <- read.csv("https://elegant-wilson-496cf4.netlify.app/practice/07-practice/data/midus.csv")

midus$sex <- factor(midus$sex)
midus$heart_self <- factor(midus$heart_self)
midus$heart_father <- factor(midus$heart_father)

midus <- na.omit(midus)

```

```{r ex1, exercise = TRUE, exercise.setup = "prepare-aovsetup"}
ggplot(data = midus, aes(x = sex, y = BMI))
  geom_violin() +
  labs(x = "Sex",
       y = "Body Mass Index (BMI)",
       title = "Exercise 1")
```

<div id="ex1-hint">

<font size="-1.5">
**Hint:** Maybe the error message is informative?
</font>
</div>


```{r ex1-solution}
ggplot(data = midus, aes(x = sex, y = BMI)) +
  geom_violin() +
  labs(x = "Sex",
       y = "Body Mass Index (BMI)",
       title = "Exercise 1")
```

```{r ex1-check}
grade_code()
```

<div data-iframe-height></div>