---
output: 
  html_document:
    includes:
      in_header: test1.html
runtime: shiny_prerendered
description: "practice app 3"
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
knitr::opts_chunk$set(echo = FALSE)
tutorial_options(exercise.eval = FALSE)
knitr::opts_chunk$set(exercise.checker = gradethis::grade_learnr)
library(palmerpenguins)
library(broom)
library(knitr)

```

```{css, echo=FALSE}
p {
  font-size: 16.5px;
  font-family: "Merriweather", sans-serif;
}

h5 {
  font-size: 16.5px;
  font-family: "Lato", sans-serif;
}

```


```{r ex1, exercise = TRUE}

```

<div id="ex1-hint">

<font size="-1.5">
y predicted by x
</font>
</div>


```{r ex1-solution}
mod3 <- lm(body_mass_g ~ sex, data = penguins)
summary(mod3)
```

```{r ex1-check}
grade_code()
```

<div data-iframe-height></div>