---
output: 
  html_document:
    includes:
      in_header: test1.html
runtime: shiny_prerendered
description: "practice app 1"
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
knitr::opts_chunk$set(echo = FALSE)
tutorial_options(exercise.eval = FALSE)
knitr::opts_chunk$set(exercise.checker = gradethis::grade_learnr)
library(palmerpenguins)

```

```{css, echo=FALSE}
p {
  font-size: 16.5px;
  font-family: "Merriweather", sans-serif;
}

h5 {
  font-size: 16.5px;
  font-family: "Lato", sans-serif;
}

```

Make sure to store your model as an object called `mod1`

```{r ex1, exercise = TRUE}
# write code to predict body mass from bill length
```

<div id="ex1-hint">

<font size="-1.5">
$ Y \sim X $ 
</font>
</div>


```{r ex1-solution}
mod1 <- lm(body_mass_g ~ bill_length_mm, data = penguins)
summary(mod1)
```

```{r ex1-check}
grade_code()
```

<div data-iframe-height></div>